// variables5.rs
//
// Execute `rustlings hint variables5` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
    let number_str = "T-H-R-E-E"; // Utiliser une variable distincte pour la chaîne de caractères
    println!("Spell a Number : {}", number_str);

    let number = 3; // Utiliser une nouvelle variable pour l'entier
    println!("Number plus two is : {}", number + 2);
}
