// enums2.rs
//
// Execute `rustlings hint enums2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

Pour définir l'énumération Message avec les différentes variantes utilisées dans la fonction main, vous pouvez utiliser les variantes suivantes :

    Move { x: i32, y: i32 } : Une variante qui contient une struct avec deux champs x et y.
    Echo(String) : Une variante qui contient une String.
    ChangeColor(u8, u8, u8) : Une variante qui contient trois valeurs u8 représentant les composantes de couleur (rouge, vert, bleu).
    Quit : Une variante qui ne contient aucune donnée. in &messages {
        message.call();
    }
}
